# DATA 601: Fall 2019

## HW1

### Due: Wed. Sep. 18 at 23:55

**Learning Objectives**

- Gain familiarity with Markdown.
- Explore collection classes in Python.
- Use intermediate level data structures and programming concepts in the context of data related problems.

_This is an individual homework assignment._ 

Please complete this homework assignment within the Jupypter notebook environment.  

#### Submission 

Your submission will be graded using a combination of auto-grading and manual grading. In order to ensure that everything goes smoothly, please follow these instructions:

- Please provide your solutions where asked; please do not alter any other parts of this notebook.
- Submit via the HW1 dropbox on D2L. Please ensure that your submitted file is named 'HW1.ipynb'.
- Do not submit any other files.

## Question 1: Markdown

#### (1 point)

Please go through the following Markdown tutorial:

  [https://www.markdowntutorial.com/](https://www.markdowntutorial.com/)
  
  In the cell below, use Markdown to typeset your name at heading level 3 and your student ID at heading level 4.

<!-- grade_id: markdown, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 2: List Processing

#### (6 points)

The questions below ask you to process data stored in lists. To focus on problem solving and to make your code more readable, *you may use built-in functions*. Please try and use comprehensions whenever possible.

1. Given two lists, write a Python function called `myprod(l1, l2)` that produces a list containing all possible pairings &mdash; as tuples &mdash; of the elements in the two lists, i.e. the Cartesian product of the two lists. The first element of each tuple should come from `l1` and the second element should come from `l2`. You can assume that the lists will not have any duplicates.  
<br>
For example, the Cartesian product of the two lists `['A', 'K', 'Q', 'J', '10', '9', '8', '7', '6', '5', '4', '3', '2']` and `['♠', '♥', '♦', '♣']` should yield a standard deck of playing cards.

<!-- grade_id: myprod, points: 0, grade: false, solution: true -->
```python
def myprod(l1, l2):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: myprod_test, points: 2, grade: true, solution: false -->
```python
'''Check that myprod produces the correct output.'''

l1 = ['A', 'K', 'Q', 'J', '10', '9', '8', '7', '6', '5', '4', '3', '2']
l2 = ['♠', '♥', '♦', '♣']
assert set(myprod(l1,l2)) == {
    ('A', '♠'), ('A', '♥'), ('A', '♦'), ('A', '♣'), \
    ('K', '♠'), ('K', '♥'), ('K', '♦'), ('K', '♣'), \
    ('Q', '♠'), ('Q', '♥'), ('Q', '♦'), ('Q', '♣'), \
    ('J', '♠'), ('J', '♥'), ('J', '♦'), ('J', '♣'), \
    ('10', '♠'), ('10', '♥'), ('10', '♦'), ('10', '♣'), \
    ('9', '♠'), ('9', '♥'), ('9', '♦'), ('9', '♣'), \
    ('8', '♠'), ('8', '♥'), ('8', '♦'), ('8', '♣'), \
    ('7', '♠'), ('7', '♥'), ('7', '♦'), ('7', '♣'), \
    ('6', '♠'), ('6', '♥'), ('6', '♦'), ('6', '♣'), \
    ('5', '♠'), ('5', '♥'), ('5', '♦'), ('5', '♣'), \
    ('4', '♠'), ('4', '♥'), ('4', '♦'), ('4', '♣'), \
    ('3', '♠'), ('3', '♥'), ('3', '♦'), ('3', '♣'), \
    ('2', '♠'), ('2', '♥'), ('2', '♦'), ('2', '♣')}
```

2. Write a Python function called `mytally( li )` that takes a list `li` and returns a dictionary whose keys are the unique entries in `li` and whose values are the counts of each of the unique entries. For example:<br><br>
      `mytally( [1, 2, 3, 3, 4, 5, 6, 6, 6, 7, 8, 9, 9] )` <br> 
      should return <br>
      `{1:1, 2:1, 3:2, 4:1, 5:1, 6:3, 7:1, 8:1, 9:2}`

<!-- grade_id: mytally, points: 0, grade: false, solution: true -->
```python
def mytally( li ):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: mytally_test, points: 2, grade: true, solution: false -->
```python
'''Check that mytally returns the correct output'''
assert mytally([1, 2, 3, 3, 4, 5, 6, 6, 6, 7, 8, 9, 9]) == \
    {1: 1, 2: 1, 3: 2, 4: 1, 5: 1, 6: 3, 7: 1, 8: 1, 9: 2}
assert mytally(list('mississippi')) == \
    {'i': 4, 'm': 1, 'p': 2, 's': 4}
import random
random.seed(a=601, version=2)
assert mytally( [random.randint(0,10) for i in range(10000)] ) == \
    {0:855, 1:895, 2:946, 3:961, 4:840, 5:941, 6:875, 7:916, 8:932, 9:910, 10:929}
```

3. Write a function called `mysplit( li )` that takes a list `li` and splits it into sublists consisting of runs of identitical elements. The returned list should be sorted in ascending order. You may assume that `li` consists of immutable and comparable objects. For exxample:<br><br>
   `mysplit( [1, 2, 3, 3, 4, 5, 6, 6, 6, 7, 8, 9, 9] )` <br> 
    should return <br>
    `[[1], [2], [3,3], [4], [5], [6,6,6], [7], [8], [9,9]]`

<!-- grade_id: mysplit, points: 0, grade: false, solution: true -->
```python
def mysplit( li ):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: mysplit_test, points: 2, grade: true, solution: false -->
```python
'''Check that mysplit produces the correct result'''
tlist = [1, 2, 3, 3, 4, 5, 6, 6, 6, 7, 8, 9, 9]
assert mysplit(tlist) == [[1], [2], [3,3], [4], [5], [6,6,6], [7], [8], [9,9]]
assert mysplit(tlist[::-1]) == [[1], [2], [3,3], [4], [5], [6,6,6], [7], [8], [9,9]]
assert mysplit(list('mississippi')) == \
       [list("iiii"), list("m"), list("pp"), list("ssss")]
```

## Question 3: Plotting Functions

#### (6 points)

Please go through the following tutorial, focusing on the first two sections.

  [https://matplotlib.org/users/pyplot_tutorial.html](https://matplotlib.org/users/pyplot_tutorial.html)
  

Use [`matplotlib.pyplot.plot`](https://matplotlib.org/users/pyplot_tutorial.html) to plot the following sequences for $ 2 \le n \le 100$. 

 a) $f_n = n^2$

 b) $f_n = \log_2 n$ 
  <br>
  (Use `math.log2(x)` to compute base 2 logarithms. You will need to `import math`)

 c) $f_n = 
    \begin{cases}
    \frac{4}{n^2 \pi^2} & \text{if } n \text{ is odd,} \\
    0 & \text{otherwise}.
    \end{cases}
    $
    (Use `math.pi` for $\pi$. You will need to `import math`)

In order to compare the relative growth rates, ***please plot within the same figure***. Use different line styles so that the sequences can be distinguished, and label your axes appropriately. Please also use logarithmic scaling on the vertical axis (`plt.yscale(‘log’)`) so that the relative magnitudes of the sequences is more apparent.


<!-- grade_id: question3, points: 6, grade: true, solution: true -->
```python
''''''
# YOUR CODE HERE
raise NotImplementedError()
```

## Question 4: Estimating a Binomial Distribution

#### (7 points)

This question asks you to empirically estimate a binomial distribution by simulating binomial trials. A high-level description of the tasks that you need to perform is provided below. You will need to think about suitable data structures and programming constructs that will accomplish the tasks. You may use [built-in](https://docs.python.org/3/library/functions.html) functions. 

1. Write a function to simulate a binomial experiment. Your function should return the number of successess in $n$ repeated trials where the probability of success for each trial is $p$. Take $n=40$ and $p=0.5$. Use a random number generator to determine if the outcome is a success or a failure. Let $\xi$ be a uniformly distributed random number in the range $[0,1)$. If $\xi < p$, then the outcome is a success, otherwise it is a failure. You can use  `math.random.random()`  to generate a uniformly distributed random `float` in the range $[0,1)$.

2. Repeat the above experiment $N$ times to determine an empirical distribution corresponding to the probability of $k$ success in $n$ trials. 
Determine two empirical distributions by taking $N=10^3$ and $N=10^6$.

3. On the same figure, plot the empirical distributions corresponding to $N=10^3$ and $N=10^6$. For comaprison, also plot the true binomial distribution for this scenario. How do the empirial distributions compare to the true distribution?

Please answer this question by inserting one ore more code cells below. Please use a Markdown cell to explain your findings.


<!-- grade_id: question4, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```